# Week 3 Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/controlflow.html) section 4 on flow control.


**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 14.1 Difference in rate per 1,000

Often in public health, we report metrics as a number per 1,000 or per 1,000,000 in population. The purpose of that is to normalize the numbers between area of larger and smaller populations.

Below, we have a function already built to calculate teh different in rate per 1,000 in two different regions.  The function normalizes them and then prints out a message describing their difference.

```python
def diff_in_rate_normalized(count_a, total_a, count_b, total_b):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B

    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population."""

    rate_a = count_a / total_a * 1000
    rate_b = count_b / total_b * 1000

    if rate_a == rate_b:
        msg = "The rate in A and the rate in B are the same ({}).".format(rate_a)
    elif rate_a > rate_b:
        msg = "The rate in A ({}) is greater than the rate in B ({}).".format(rate_a, rate_b)
    else:
        msg = "The rate in A ({}) is less than the rate in B ({})".format(rate_a, rate_b)

    return msg
```

```python
rate_covid_testing = diff_in_rate_normalized(1,1000,15,10000)
print(rate_covid_testing)
```

Output:
```
The rate in A (1.0) is less than the rate in B (1.5)
```

```python
diff_in_rate_normalized(3,1000,30,10000)
```

Output:
```
'The rate in A and the rate in B are the same (3.0).'
```

### 14.2 Trimming outliers

While it works well in most cases, in practice, this may not truly make numbers comparable between extremely large population centers (e.g. New York City at 8.5 million) and very small rural areas (e.g. Meeteese, WY at 459 people).

Let's take that function and make some adjustments.  If the total population of either A or B is more than 100 times larger than the other, then we want to return a different message.  That is, if the populations are more than two orders of magnitude different, then we shouldn't try to compare them.

```python

```

```python
def diff_in_rate_normalized(count_a, total_a, count_b, total_b):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B

    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population.

    If total_a / total_b > 100 OR total_b / total_a > 100 then we'll return a message saying the two can't be compared.
    "The total populations in A and B are so different that they can't be compared."
    """
     # Calculate rates
    rate_a = count_a / total_a * 1000
    rate_b = count_b / total_b * 1000
     # Compare rates and create message
    if rate_a == rate_b:
        msg = "The rate in A and the rate in B are the same ({}).".format(rate_a)
    elif rate_a > rate_b:
        msg = "The rate in A ({}) is greater than the rate in B ({}).".format(rate_a, rate_b)
    else:
        msg = "The rate in A ({}) is less than the rate in B ({}).".format(rate_a, rate_b)

    return msg


```

```python
diff_in_rate_normalized(5, 459, 30, 8500)
```

Output:
```
'The rate in A (10.893246187363834) is greater than the rate in B (3.5294117647058827).'
```

### 14.3 LACE Score
There is a simple readmission index called the LACE Score: https://www.hindawi.com/journals/bmri/2015/169870/tab1/

Use this documentation to create a function that can compute LACE scores based on the 4 input parameters.  Below is the function signature and documentation that you start with.

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def LACE(length_of_stay, acute_flag, charlson, ed_visits):
    """(int, bool, int, int) -> int
    This function computes the LACE score for this patient based on
    the logic from https://www.hindawi.com/journals/bmri/2015/169870/tab1/.

    >>> LACE(4, False, 1, 0)
    5

    >>> LACE(4, True, 4, 7)
    16
    """
    # Initialize the score to 0
    score = 0

    # Scoring criteria for length of stay
    length_of_stay_scores = [
        (14, 7),   # Long stay (14 days or more)
        (7, 5),    # Medium-long stay (7 to 13 days)
        (4, 4),    # Medium stay (4 to 6 days)
        (3, 3),    # Short stay (3 days)
        (2, 2),    # Very short stay (2 days)
        (1, 1)     # One-day stay (1 day)
    ]

    # Calculate score based on the length of stay using a loop
    for threshold, points in length_of_stay_scores:
        if length_of_stay >= threshold:
            score += points
            break  # Exit the loop after finding the score

    # Add score if the patient was admitted as an acute case
    if acute_flag:
        score += 3  # Acute care adds 3 points

    # Scoring criteria for Charlson comorbidity index
    charlson_scores = [
        (4, 5),  # High comorbidity score (4 or more)
        (3, 3),  # Moderate comorbidity score (3)
        (2, 2),  # Mild comorbidity score (2)
        (1, 1)   # Minimal comorbidity score (1)
    ]

    # Calculate score based on Charlson comorbidity index using a loop
    for threshold, points in charlson_scores:
        if charlson >= threshold:
            score += points
            break  # Exit the loop after finding the score

    # Scoring criteria for emergency department visits
    ed_visit_scores = [
        (4, 4),  # Frequent ED visits (4 or more)
        (3, 3),  # Three ED visits get 3 points
        (2, 2),  # Two ED visits get 2 points
        (1, 1)   # One ED visit gets 1 point
    ]

    # Calculate score based on the number of emergency department visits using a loop
    for threshold, points in ed_visit_scores:
        if ed_visits >= threshold:
            score += points
            break  # Exit the loop after finding the score

    # Return the final LACE score
    return score


```

<!-- grade_id: bmi-test, points: 10, grade: true, solution: false -->
```python
assert LACE(4, False, 1, 0) == 5, "case 1 failed"
assert LACE(4, True, 4, 7) == 16, "case 2 failed"

print("All assertions passed!")


```

Output:
```
All assertions passed!
```

```python

```

### 14.4 Care Management Criteria

Care managers use LACE as part of the criteria for assigning a care coodinator to a patient who has been recently discharged. If the score is above 10, then a care coordinator will be assigned. The other criteria they use is if the patient has been discharged with a diagnosis of CHF or COPD.  If the diagnosis field has CHF or COPD in it, then the patient will have a care coordinator assigned.

For this exercise, write another function that takes the same inputs as LACE() plus another diagnosis parameter, and return True or False depending on if the paient needs a care coordinator.

**NOTE** Pay attention to the fact that the order of parameters in this function definition are not the same as the order they were in the LACE score.  

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def assign_care_coordinator(diagnosis_cd, ed_visits, length_of_stay, acute_flag, charlson):
    """ (str, int, int, bool, int) -> bool
    Care managers use LACE as part of the criteria for assigning a care coordinator to a
    patient who has been recently discharged. If the score is above 10, then a care
    coordinator will be assigned. The other criteria they use is if the patient has been
    discharged with a diagnosis of CHF or COPD. If the diagnosis field has CHF or COPD in
    it, then the patient will have a care coordinator assigned.

    >>> assign_care_coordinator('None', 0, 4, False, 1)
    False

    >>> assign_care_coordinator('CHF', 0, 4, False, 1)
    True

    >>> assign_care_coordinator('COPD', 0, 4, False, 1)
    True

    >>> assign_care_coordinator('None', 7, 4, True, 4)
    True

    >>> assign_care_coordinator('CHF', 7, 4, True, 4)
    True
    """

    # Check for specific diagnoses
    if diagnosis_cd in ['CHF', 'COPD']:
        return True  # Assign care coordinator

    # Initialize the score to 0
    score = 0

    # Scoring criteria for length of stay
    length_of_stay_scores = [
        (14, 7),  # Long stay (14 days or more)
        (7, 5),   # Medium-long stay (7 to 13 days)
        (4, 4),   # Medium stay (4 to 6 days)
        (3, 3),   # Short stay (3 days)
        (2, 2),   # Very short stay (2 days)
        (1, 1)    # One-day stay (1 day)
    ]

    # Calculate score based on length of stay using a loop
    for threshold, points in length_of_stay_scores:
        if length_of_stay >= threshold:
            score += points
            break  # Exit the loop after finding the score

    # Add score if the patient was admitted as an acute case
    if acute_flag:
        score += 3  # Acute care adds 3 points

    # Scoring criteria for Charlson comorbidity index
    charlson_scores = [
        (4, 5),  # High comorbidity score (4 or more)
        (3, 3),  # Moderate comorbidity score (3)
        (2, 2),  # Mild comorbidity score (2)
        (1, 1)   # Minimal comorbidity score (1)
    ]

    # Calculate score based on Charlson comorbidity index using a loop
    for threshold, points in charlson_scores:
        if charlson >= threshold:
            score += points
            break  # Exit the loop after finding the score

    # Scoring criteria for emergency department visits
    ed_visit_scores = [
        (4, 4),  # Frequent ED visits (4 or more)
        (3, 3),  # Three ED visits get 3 points
        (2, 2),  # Two ED visits get 2 points
        (1, 1)   # One ED visit gets 1 point
    ]

    # Calculate score based on the number of emergency department visits using a loop
    for threshold, points in ed_visit_scores:
        if ed_visits >= threshold:
            score += points
            break  # Exit the loop after finding the score

    # Assign a care coordinator if the total score exceeds 10
    return score > 10


```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert assign_care_coordinator('None', 0, 4, False, 1) == False
assert assign_care_coordinator('CHF', 0, 4, False, 1) == True
assert assign_care_coordinator('COPD', 0, 4, False, 1) == True
assert assign_care_coordinator('None', 7, 4, True, 4) == True
assert assign_care_coordinator('CHF', 7, 4, True, 4) == True
```

```python

```

### 14.5 qCSI COVID-19 Severity Index

See: https://www.mdcalc.com/quick-covid-19-severity-index-qcsi#evidence

Calculate the total risk score as per the point values assigned to respiratory rate, pulse oximetry, and O2 flow rate.  Then calculate and return the Risk Level.

In addition to the rules provided at the link above, also add the following checks for valid values:
* If `respiratory_rate <= 0` then return _invalid respiratory rate_
* If `pulse_ox <= 0` then return _invalid pulse ox_
* If `pulse_ox > 100` then return _invalid pulse ox_
* If `os_flow <= 0` then return _invalid O2 flow rate_

```python
def qcsi(respiratory_rate, pulse_ox, o2_flow):
    """(int, int, int) -> str
    * respiratory_rate is an integer value
    * pulse_ox is an integer value (e.g. 30 means 30%)
    * o2_flow is an integer value (liters of oxygen flow per minute)
    This function calculates a QCSI score based on the input parameters and determines a risk level.
    """

    # Initialize scores
    score_1 = score_2 = score_3 = 0

    # Scoring criteria for respiratory rate
    respiratory_rate_scores = [
        (28, 2),  # High respiratory rate
        (23, 1),  # Intermediate respiratory rate
        (22, 0)   # Low or normal respiratory rate
    ]

    # Calculate score for respiratory rate using a loop
    for threshold, points in respiratory_rate_scores:
        if respiratory_rate > threshold:
            score_1 = points
            break
        elif respiratory_rate <= threshold:
            score_1 = points

    # Scoring criteria for pulse oxygen level
    pulse_ox_scores = [
        (88, 5),  # Very low oxygen saturation
        (92, 2),  # Moderate oxygen saturation
        (100, 0)  # Normal oxygen saturation
    ]

    # Calculate score for pulse oxygen level using a loop
    for threshold, points in pulse_ox_scores:
        if pulse_ox <= threshold:
            score_2 = points
            break

    # Scoring criteria for oxygen flow rate
    o2_flow_scores = [
        (6, 5),  # High oxygen flow rate
        (4, 4),  # Intermediate flow rate
        (2, 0)   # Low or no oxygen flow
    ]

    # Calculate score for oxygen flow rate using a loop
    for threshold, points in o2_flow_scores:
        if o2_flow >= threshold:
            score_3 = points
            break

    # Total score is the sum of all three scores
    total_score = score_1 + score_2 + score_3

    # Validation checks for input values
    if respiratory_rate <= 0:
        return "invalid respiratory rate"  # Respiratory rate can't be 0 or less
    elif pulse_ox <= 0 or pulse_ox > 100:
        return "invalid pulse ox"  # Pulse ox must be between 1 and 100
    elif o2_flow <= 0:
        return "invalid O2 flow rate"  # O2 flow rate can't be 0 or negative

    # Risk assessment based on the calculated score
    if total_score <= 3:
        risk_level = "low"  # Low risk
    elif 4 <= total_score < 6:
        risk_level = "low-intermediate"  # Low-intermediate risk
    elif 7 <= total_score < 9:
        risk_level = "high-intermediate"  # High-intermediate risk
    elif 10 <= total_score < 12:
        risk_level = "high"  # High risk

    return risk_level  # Return the risk level based on the score


```

```python
assert (qcsi(29, 95, 1) == 'low')
assert (qcsi(20, 93, 1) == 'low')
assert (qcsi(29, 88, 1) == 'high-intermediate')
assert (qcsi(29, 88, 4) == 'high')
assert (qcsi(30, 90, 1) == 'low-intermediate')
assert (qcsi(28, 92, -1) == 'invalid O2 flow rate')
assert (qcsi(22, 0, 4) == 'invalid pulse ox')
assert (qcsi(0, 97, 2) == 'invalid respiratory rate')
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week03/week03_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**



---