# Week 4 Exercises

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 13 through 15 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 20.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names.

```
sort_names([
  ['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']])

['Paul Boal', 'Donald Duck', 'Kermit the Frog']
```

**Tip:** If you have a list of lists and want to sort it, Python can do that. It will simply sort using the first item in each of the lists inside your list.

```python
names = [['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']]
```

```python
def sort_names(names):
    # Join each list of names into a single string
    full_names = [' '.join(name) for name in names]
    # Sort the full names alphabetically
    full_names.sort()
    return full_names

```

```python
names_list = [
    ['Paul', 'Boal'],
    ['Kermit', 'the', 'Frog'],
    ['Donald', 'Duck']
]
sorted_names = sort_names(names_list)
print(sorted_names)

```

Output:
```
['Donald Duck', 'Kermit the Frog', 'Paul Boal']
```

### 20.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Elliie'],
 ['Lester-Boal', 'Ada'],
 ...
```

```python

```

```python
def dict_to_list(families):
    # Initialize an empty list to hold the output
    output = []

    # Iterate through the dictionary items
    for family, names in families.items():
        # Create a sublist for each name in the family
        for name in names:
            output.append([family, name])

    return output


```

```python

families = {
    'Boal': ['Paul', 'Anny', 'Jim'],
    'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
    'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

# Convert the dictionary to a list
output = dict_to_list(families)
print(output)

```

Output:
```
[['Boal', 'Paul'], ['Boal', 'Anny'], ['Boal', 'Jim'], ['Lester-Boal', 'Ellie'], ['Lester-Boal', 'Ada'], ['Lester-Boal', 'Teddy'], ['Lester', 'Sarahlynn'], ['Lester', 'Jessica'], ['Lester', 'Grace'], ['Lester', 'Carolynn']]
```

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():
    names = [['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    # Initialize an empty dictionary to hold the families
    families = {}

    # Loop through each entry in the names list
    for family, name in names:
        # If the family is not already a key in the dictionary, add it
        if family not in families:
            families[family] = []
        # Append the name to the list of names for that family
        families[family].append(name)

    return families


```

```python
output = list_to_dict()

```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert output == answer
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):

    conclusions = []
   # Loop through each patient record
    for patient in patients:
        name = patient[0]
        diagnosis = patient[1]
        length_of_stay = patient[2]

        # Get the average length of stay for the diagnosis
        average_length = avg_los.get(diagnosis)

        # Compare the patient's length of stay with the average
        if average_length is not None:  # Check if the diagnosis exists in the dictionary
            if length_of_stay < average_length:
                conclusions.append([name, 'too short'])
            elif length_of_stay > average_length:
                conclusions.append([name, 'too long'])
            else:
                conclusions.append([name, 'just right'])

    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
[['Boal', 'too short'],
 ['Boal', 'too long'],
 ['Jones', 'just right'],
 ['Ashbury', 'too short']]
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week04/week04_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

